# <div style="background-color:rgba(204, 229, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(0, 76, 153, 1);">**PHYS 121 Pre-Lab #4**<span style="color:red"> $\to$ (5 possible marks)</span> </span></div> 
# Electric & Hydraulic Circuits – Week 1

***
## Learning Objectives:
* <b><span style="color:rgba(0, 153, 76, 1);"> Perform simple operations on DataFrames and arrays in Python.</span></b>

***

Over the next two weeks, you will work with both electric and hydraulic circuits.  You will see that concepts developed to analyze electrical circuits can also be applied to their hydraulic analogues.  In addition, you will attempt to uncover differences in the behaviour of the two systems and thus identify limitations of the analogy.) 

In the process, you will be collecting and manipulating data.  As before, a lot of code will be provided for you, but there will be instances where you will be asked to input your own lines of code.  This Pre-Lab is intended to provide a brief review of the kinds of operations you can use in Python and an introduction to how these operations can be applied to DataFrames and/or arrays. 

***
## Autograding:
The PHYS 121 Pre-lab assignments and Labs will make use of some autograding.  To make the autograding work, the two cells below needs to be executed.  The first one installs the required packages and the second imports required packages/modules.  If 'PHYS121.Installer()' reports that some functions have been installed, the user should restart the "kernel" and then exectute the PHYS121.Installer() cell a second time.  The second time the installer function is run, it should report that **"All packages already installed. Please proceed"**.

If necessary, the kernel can be restarted by selecting **Kernel** $\to$ **Restart Kernel** from the menu options at to the top of the page.  Here is a <a href = "https://cmps-people.ok.ubc.ca/jbobowsk/PHYS_121_Lab/images/restartKernel.gif">GIF</a> showing how to restart the kernel.

The 'PHYS121.Installer()' command requires the file 'PHYS121.py', which you should see included in the list of files along the left-hand side of the screen.

```python
# Import PHYS121.py and then run the installer function.
import PHYS121
PHYS121.Installer()
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("PHYS 121 - Pre-Lab 4.ipynb")
```

***
## Import Modules:
Execute the cell below to import a number of useful pre-built Python modules that will be used in the PHYS 121 pre-labs and labs.

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
from matplotlib.pyplot import cm # used to generate a sequence of colours for plotting
from scipy.optimize import curve_fit
from IPython.display import HTML as html_print
from IPython.display import display, Markdown, Latex
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 1 - Basic Operations</span></div>

Recall the basic arithmetic operations in Python:

```python
# Remember: to execute the code cell, select it and press shift + enter

# Addition / Subtraction
a = 4 + 5
b = 2 - 3

a, b
```

```python
# Division / Miltiplication
x = 10.2 * -4.0
y = 2 / 10

x, y
```

```python
# Exponentiation
s = 12**2

s
```

With NumPy, we aquire a larger set of operations as well as a host of constants to work with. The notation is fairly straight-forward.  Execture the cells below to see how these simple mathematical operations can be implemented.

```python
# Pi and Euler's number e can be obtained as follows:
np.pi, np.e
```

```python
# Natural logarithm
a = np.log(1)
b = np.log(np.e)

a, b
```

```python
# Trigonometric functions
c = np.cos(0)
d = np.sin(np.pi)
e = np.tan(np.pi / 2)

c, d, e
```

Be cautious when using these functions: they expect the input to be in units of radians. If you have a measurement in degrees, you can convert it to radians in the following way:

```python
# Convert 180 degrees to radian
theta = np.radians(180)
f = np.sin(theta)

theta, f 
```

You should also notice the rounding errors in the above. We should have $\sin\pi = 0$ and $\tan(\pi / 2) = \infty$, but instead we get something like $1.22\times10^{-16}$ and $1.63\times10^{16}$. These kinds of errors occur all the time in floating-point arithmetic, so we sometimes have to be careful about what we're doing. If you're interested, [this article](https://docs.python.org/3/tutorial/floatingpoint.html) covers the phenomenon in more detail, but you shouldn't have to worry too much about it for this lab course. 

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 2 - Operating on Arrays and DataFrames</span></div>

Now, suppose we wanted to perform an operation on a set of data. All of the operations discussed above can be performed on entire NumPy arrays or columns from DataFrames. For instance, consider pairwise addition of two datasets:

```python
# When we add two arrays together, the output is another array
array1 = np.array([1, 2, 3])
array2 = np.array([4, 5, 0])

array3 = array1 + array2

array3
```

First, use the Pandas package (imported above as 'pd') to define a dataframe consisting of two columns.

```python
# Define a dataframe called 'df' with two columns
df = pd.DataFrame({"c1" : [1, 2, 3],
                   "c2" : [4, 5, 0],
                  })

# Show the dataframe 'df'
df
```

When we add two columns from a dataframe together, the output is an object called a "series", which can be used to create an additional column.

```python
# Create a third column from the sum of the first two 
df["c3"] = df["c1"] + df["c2"]

# Show the updated dataframe 'df' which now has three columns
df
```

Functions from the NumPy package can be applied to both DataFrames and NumPy arrays.

```python
# Trig functions act on each element individually
array4 = np.sin(array1)

# Display array4
array4
```

Multiplication of arrays or DataFrame columns is performed pairwise.  The output is another array/column of the same size as the original arrays/columns.

```python
# Create a fourth column in the 'df' dataframe which is formed from the pairwise product of the elements in the first two columns. 
df["c4"] = df["c2"] * df["c1"]

# Display the dataframe which now has a total of four columns
df
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 3 - Snell's Law Data Analysis</span></div>

The following section will guide you through the analysis of some experimental data, giving you an opportunity to practice performing operations on DataFrames. First, a bit of background on optics...

**Snell's law** describes the change in direction of light as it passes from one medium to another, called **refraction**. Refer to Fig 1. The incoming light strikes the interface at an angle of $\theta_1$ from perpendicular (the **angle of incidence**), and the outgoing light leaves at an angle $\theta_2$ (the **angle of refraction**). The relationship between these angles can be described in terms of the **refractive index** of the materials in question, denoted $n_1$ and $n_2$. Then Snell's law states that

\begin{align}
\frac{\sin\theta_1}{\sin\theta_2} = \frac{n_2}{n_1}
\label{eq:snell} \tag{1}
\end{align}

In the figure below $\sin\theta_1 > \sin\theta_2$, such that medium two must have a larger refractive index than medium one (i.e. $n_2 > n_1$).

<p>
<center>
<img src="https://cmps-people.ok.ubc.ca/jbobowsk/PHYS_121_Lab/Lab4/images/snell.svg" alt="image info" style="width: 35%; height: auto;" /><br>
<b>Fig. 1: The refraction of light passing from one medium to another. The incoming light strikes the interface at an angle of $\theta_1$ from perpendicular, and the outgoing light leaves at an angle $\theta_2$. The relationship between these angles and the refractive index of the materials is given by Snell's law. </b></center>
</p>

If the first medium is air, these angles can be measured by shining a beam towards a hemispherical dish made of (or filled with) the second medium, as in Fig. 2. If the beam strikes the centre of the dish, no matter what the angle of refraction $\theta_2$ is, the beam will exit the dish perpendicular to its surface, and will not be refracted.

<p>
<center>
<img src="https://cmps-people.ok.ubc.ca/jbobowsk/PHYS_121_Lab/Lab4/images/snell2.svg" alt="image info" style="width: 45%; height: auto;" /><br>
<b>Fig. 2: The angles $\theta_1$ and $\theta_2$ can be measured using a hemispherical dish. </b></center>
</p>

The refractivity of a meterial can tell you about its composition. For example, the refractive index of seawater will increase as its salinity (or salt concentration) is increased. The code cell below imports some generated "experimental" data on a sample of seawater held at $20^\circ\rm C$. We are going to analyze this data to determine the index of refraction of the seawater, which can then be used to find the salinity of the water. 

The data is contained in a file called 'SampleData.csv'.  This file should already appear among the files listed along the left-had side of the screen.

```python
# Import the contents of 'SampleData.csv' into a dataframe called 'data'
data = pd.read_csv("SampleData.csv")

# Display the contents of the 'data' dataframe.
data
```

### <span style="color:purple">The 'grader.check(...)' statements in this Pre-Lab only check that answers have been entered in the correct format.  They do **NOT** check that you've entered the correct answer.</span> 

***
**<span style="color:blue">Question 3.1</span>**  **<span style="color:red">(1 mark)</span>**

First, we will calculate the sine of the angles in the 'theta1' and 'theta2' columns.  These values will be added as new columns to our DataFrame called 'data'. The code cell below does this for 'theta1' column. Use the empty code cell that follows to repeat the process for the 'theta2' column. 

*** Please do not change any of the variable names in the cells below ***

```python
# Remember that we must also convert the angles to radians.
data["sin1"] = np.sin(np.radians(data["theta1"]))
data
```

***
**<span style="color:blue">Answer 3.1:</span>**

```python
# Repeat the process above for the angle of refraction 'theta2'
data["sin2"] = ...
data
```

```python
grader.check("Q3.1")
```

***
**<span style="color:blue">Question 3.2</span>**  **<span style="color:red">(1 mark)</span>**

Now, if we calculate the ratio of these angles, we can apply Snell's law to determine the index of refraction of the water using Eq. (1). Using the code cell below, create a new column in the 'data' DataFrame constructed from the ratio of the 'sin1' and 'sin2' columns (i.e. sin1/sin2). 

*** Please do not change any of the variable names in the cells below ***

***
**<span style="color:blue">Answer 3.2:</span>**

```python
data["ratio"] = ...
data
```

```python
grader.check("Q3.2")
```

***
**<span style="color:blue">Question 3.3</span>**  **<span style="color:red">(1 mark)</span>**

The refractive index of air is $1.000227$. Using $n_1 = 1.000227$ and Eq. (1), create a new column in the 'data' dataframe with the exerpimentally-determined value of the refractive index $n_2$ of the seawater.

*** Please do not change any of the variable names in the cells below ***

***
**<span style="color:blue">Answer 3.3:</span>**

```python
data["n2"] = ...
data
```

```python
grader.check("Q3.3")
```

We can take the average value and standard error in the mean for this data directly with Pandas. The syntax for this is 

```python
df["column"].mean()
df["column"].sem()
```

for the mean and standard error, respectively. In the syntax above, 'df' is the name of the dataframe ('data', in our case) and "column" needs to be replaced with the name of the column that you are analyzing.

***
**<span style="color:blue">Question 3.4</span>**  **<span style="color:red">(1 mark)</span>**

Use the code cell below to calculate the mean and standard error for the refractive index of our seawater.

*** Please do not change any of the variable names in the cells below ***

***
**<span style="color:blue">Answer 3.4:</span>**

```python
avg = ...
sem = ...

print(f"n_2 = {avg:.4f} +/- {sem:.4f}")
```

```python
grader.check("Q3.4")
```

In a 1995 paper, [Quan & Fry](https://opg.optica.org/ao/fulltext.cfm?uri=ao-34-18-3477&id=45728) provide an empirical equation for the refractive index of seawater as a function of salinity, temperature, and the wavelength of light. For a wavelength of light equal to $600\rm\ nm$ and a temperature of $10\rm ^\circ C$, Quan & Fry's equation becomes

\begin{align}
S(n) = 5312.08 n  - 7082.90
\label{eq:salinity} \tag{2}
\end{align}

where $S$ is the salinity in grams per litre.  The uncertainty in the salinity is given by

\begin{align}
\sigma_S = 5312.08 \sigma_n
\label{eq:salinityerr} \tag{3}
\end{align}

where $\sigma_n$ is the uncertainty in the refractive index.

***
**<span style="color:blue">Question 3.5</span>**  **<span style="color:red">(1 mark)</span>**

Using the code block below, calculate the salinity of the sample and its uncertainty. 

*** Please do not change any of the variable names in the cells below ***

***
**<span style="color:blue">Answer 3.5:</span>** 

```python
S = ...
errS = ...

# Print the results.  The "{:.2f}".format() synatx is used to display the results with two decimals of precision.
print('S =', "{:.2f}".format(S), '+/-', "{:.2f}".format(errS), 'g/L')
```

```python
grader.check("Q3.5")
```

And that's all there is to it! Now that you understand the basics, you'll be able to apply them in this week's lab.

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 4 - Feedback and Submission</span></div>

<!-- BEGIN QUESTION -->

**<span style="color:blue">Question 4.1</span>**  

We welcome your feedback on the PHYS 121 pre-labs!  Please feel free to include any comments you have about this pre-lab in the cell below.  Your comments will be taken into consideration when revising/improving the PHYS 121 labs and pre-labs.  You can suggest improvements, point out anything that was unclear, comment on the strengths and weaknesses of the pre-lab, ...

This question is optional and will have no impact on your pre-lab or lab grade.

***
**<span style="color:blue">Answer 4.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

***

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Once you've completed this notebook:
- Save your work.
- Run 'grader.export()' to generate a .zip file containing all of the materials that you will submit.
- Download the generated .zip file.
- Upload the .zip file to the PHYS 121 Lab Canvas gradebook.

Here is a <a href = "https://cmps-people.ok.ubc.ca/jbobowsk/phys231/Python/images/Submission.gif">GIF</a> showing how these steps are completed.  Once your completed notebook has been uploaded to the Canvas gradebook, you're done!

```python
# Save your notebook first, then run this cell (place your cursor in the cell and then hit Shift + Enter) to export 
# your submission.  Uploaded the .zip file that is generated to the gradebook of the Canvas shell for the PHYS 121 lab.
grader.export()
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 5 - Playground (optional)</span></div>

Feel free to add as many cells as you like below and use them as a playground for further independent investigations.  These cells won't be graded, so feel free to use them in any way that you like.  For example, you could compare Gaussian distributions with different standard deviations and/or means. 

```python
# Here's an empty code cell that you can use.

```

```python
# Here's another empty code cell that you can use.

```

```python
# Here's yet another empty code cell that you can use.  
# If you need more, you can add cells using the '+' icon in the menu bar at to the top of the screen.

```

***
Last update: February 1, 2023

Tests for Q3.1 (points: None)
```
>>> isinstance(data["sin2"], pd.core.series.Series)
True
```
failure_message: Your data is not the correct type (Pandas Series). Refer to the syntax above used for data["sin1"] to define data["sin2"].
```
>>> len(data["sin2"]) == 17
True
```
failure_message: Your data is not the correct length. Refer to the syntax above used for data["sin1"] to define data["sin2"].


Tests for Q3.2 (points: None)
```
>>> isinstance(data["ratio"], pd.core.series.Series)
True
```
failure_message: Your data is not the correct type (Pandas Series).
```
>>> len(data["ratio"]) == 17
True
```
failure_message: Your data is not the correct length.


Tests for Q3.3 (points: None)
```
>>> isinstance(data["n2"], pd.core.series.Series)
True
```
failure_message: Your data is not the correct type (pandas Series).
```
>>> len(data["n2"]) == 17
True
```
failure_message: Your data is not the correct length.


Tests for Q3.4 (points: None)
```
>>> isinstance(avg, float)
True
```
failure_message: Your "avg" variable is not the correct type (numpy float). Make sure you are using the built in .mean() function listed above.
```
>>> isinstance(sem, float)
True
```
failure_message: Your "sem" variable is not the correct type (numpy float). Make sure you are using the built in .sem() function listed above.


Tests for Q3.5 (points: None)
```
>>> isinstance(S, float)
True
```
failure_message: Your "S" variable is not the correct type (numpy float).
```
>>> isinstance(errS, float)
True
```
failure_message: Your "errS" variable is not the correct type (numpy float).
